---
title: Introduction
description: "1. Understand the objectives of biostatistics and Health Services Research\n2. Explore data types, distributions, and the basic ways of displaying data"
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
